## Assignment 2

#### Feel free to add cells to define your helper functions.
#### Please remove "raise NotImplementedError()"

Student 1 Name: Tamara Zubatiy A11904774

Student 2 Name: Grant Nelson

Tamara Zubatiy and Grant Nelson attest that this assignment was done by them two and reflects their original work and based on their understanding of the concepts. Both students have equally contributed to the solution of this assignment.

```python
import numpy as np
import sklearn 
import time
import random
from sklearn import datasets
from sklearn import metrics
from sklearn.neighbors import KNeighborsClassifier
from sklearn.linear_model import Perceptron
from sklearn.svm import SVC
from sklearn.discriminant_analysis import LinearDiscriminantAnalysis
import os
    
```

<!-- grade_id: cell-382d24a442ac36ea, points: 0, grade: false, solution: true -->
```python
def sentimentClassKNN(trainFile, testFile, k, distanceMeasure):
    # Open and read the trainFile and testFile
    var= open(trainFile, 'r')
    tvar= open(testFile, 'r')

    
    #random shuffle trainFile before splitting labels
    train = []
    test = []

    train= var.readlines()
    test= tvar.readlines()
    
    var.close()
    tvar.close()
        
    random.shuffle(train)  
    random.shuffle(test)
    
    trainLabels=[]
    trainData=[]
    testLabels=[]
    testData=[]
        
    # Split label and vector
    for elements in train:
        trainLabels.append(float(elements.split()[0]))
        trainData.append([float(i) for i in elements.split()[1:]])
        
    for elements in test:
        testLabels.append(float(elements.split()[0]))
        testData.append([float(i) for i in elements.split()[1:]])
        
    # Create a KNN model
   
    #k-N-N
    model=KNeighborsClassifier(n_neighbors = k, metric= distanceMeasure) #accuracy 81.85% with your k value
    
    

    # Fit the training set into the model
    model.fit(trainData, trainLabels)
    
    # Calculate and return the prediction accuracy
    accuracy = model.score(testData, testLabels)
    print(accuracy)
    return accuracy
    
    
   
```

<!-- grade_id: cell-5efea8601cd9b01e, points: 10, grade: true, solution: false -->
```python
knn_accu = sentimentClassKNN('train_data.txt', 'test_data.txt', 100, "euclidean")
assert (0.8<knn_accu)
```

Output:
```
0.8185
```

<!-- grade_id: cell-cdb0a02d14ecdb4f, points: 0, grade: false, solution: true -->
```python
def sentimentClassPerceptron(trainFile, testFile, n):
    # Open and read the trainFile and testFile
    var= open(trainFile, 'r')
    tvar= open(testFile, 'r')

    
    #random shuffle trainFile before splitting labels
    train = []
    test = []

    train= var.readlines()
    test= tvar.readlines()
    
    var.close()
    tvar.close()
        
    random.shuffle(train)  
    random.shuffle(test)
    
    trainLabels=[]
    trainData=[]
    testLabels=[]
    testData=[]
        
    # Split label and vector
    for elements in train:
        trainLabels.append(float(elements.split()[0]))
        trainData.append([float(i) for i in elements.split()[1:]])
        
    for elements in test:
        testLabels.append(float(elements.split()[0]))
        testData.append([float(i) for i in elements.split()[1:]])
     #make the model
   
    model=sklearn.linear_model.Perceptron(n_iter=n)
     # Fit the training set into the model
    model.fit(trainData, trainLabels)
    
    # Calculate and return the prediction accuracy
    accuracy = model.score(testData, testLabels)
    print(accuracy)
    return accuracy
    
    
    
    raise NotImplementedError()
    
```

<!-- grade_id: cell-d53f98d0f4e7704b, points: 10, grade: true, solution: false -->
```python
per_accu = sentimentClassPerceptron('train_data.txt', 'test_data.txt', 100)
assert (0.8 < per_accu)
```

Output:
```
0.8179
```

<!-- grade_id: cell-b5ca32a7c5be53ed, points: 0, grade: false, solution: true -->
```python
def sentimentClassMyOwn(trainFile, testFile):

    # Open and read the trainFile and testFile
    var= open(trainFile, 'r')
    tvar= open(testFile, 'r')

    
    #random shuffle trainFile before splitting labels
    train = []
    test = []

    train= var.readlines()
    test= tvar.readlines()
    
    var.close()
    tvar.close()
        
    random.shuffle(train)  
    random.shuffle(test)
    
    trainLabels=[]
    trainData=[]
    testLabels=[]
    testData=[]
        
    # Split label and vector
    for elements in train:
        trainLabels.append(float(elements.split()[0]))
        trainData.append([float(i) for i in elements.split()[1:]])
        
    for elements in test:
        testLabels.append(float(elements.split()[0]))
        testData.append([float(i) for i in elements.split()[1:]])
     
    
    #make the first model Logistic Regression
   
    model1=sklearn.linear_model.LogisticRegression()
     # Fit the training set into the model
    model1.fit(trainData, trainLabels) #TOOK 4.2 SECONDS TO RUN WITH ACCURACY 86.76%
    
    
     #make the second model SVM but use the support vector classifier class 
   
    model2=sklearn.svm.SVC()
     # Fit the training set into the model
    model2.fit(trainData, trainLabels) #TOOK 5 MINUTES TO RUN WITH ACCURACY 87.19%

    #make the third model Fisher Linear Discriminant Analysis  
    model3=sklearn.discriminant_analysis.LinearDiscriminantAnalysis()
     # Fit the training set into the model 
    model3.fit(trainData, trainLabels) #TOOK 2.69 SECONDS TO RUN WITH ACCUURACY 86.74%

    # Calculate and return the prediction accuracy
    accu1 = model1.score(testData, testLabels)
    accu2 = model2.score(testData, testLabels)
    accu3 = model3.score(testData, testLabels)
    print(accu1)
    print(accu2)
    print(accu3)
    
    
    
    return accu1, accu2, accu3
    
    
```

<!-- grade_id: cell-5d9c51d89449cc94, points: 30, grade: true, solution: false -->
```python
accu1, accu2, accu3 = sentimentClassMyOwn('train_data.txt', 'test_data.txt')
assert (0.8 < accu1)
assert (0.8 < accu2)
assert (0.8 < accu3)
```

Output:
```
0.8676
0.8719
0.8674
```

<!-- grade_id: cell-245e801b01c4fe2a, points: 0, grade: true, solution: false -->
```python
### Check to make sure you include your report.
assert os.path.isfile('Assignment2_Report.pdf')
```

```python

```